# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!

MYTABLE='paulboal' + '_data'



# Then, this code will automatically setup your 
import os
from sqlalchemy import create_engine
import pandas as pd

host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + ':' + 
                     str(port) + '/' + 
                     database)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'pboal')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'

### BEGIN SOLUTION
df = pd.read_csv(url)
df.to_sql(MYTABLE, conn, if_exists='replace')
### END SOLUTION
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,20))
assert(list(dxyz.columns) == ['index', 'Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
```

```python
df.head()
```

Output:
```
Provider_Name         NPI    CCN  \
0                          Sutter Bay Hospitals  1659439834  50008   
1  PRIME HEALTHCARE SERVICES - GARDEN GROVE LLC  1659538858  50230   
2                        ST MARY MEDICAL CENTER  1669456299  50300   
3                     MADERA COMMUNITY HOSPITAL  1669673646  50568   
4                  Temecula Valley Hospital Inc  1679816201  50775   

  Medicaid_EP_Hospital_Type           Street_Address           City  \
0      Acute Care Hospitals      CASTRO & DUBOCE STS  SAN FRANCISCO   
1      Acute Care Hospitals  12601 GARDEN GROVE BLVD   GARDEN GROVE   
2      Acute Care Hospitals      18300 US HIGHWAY 18   APPLE VALLEY   
3      Acute Care Hospitals        1250 E ALMOND AVE         MADERA   
4      Acute Care Hospitals      31700 TEMECULA PKWY       TEMECULA   

           County State  Zip_Code  Payment_Year_Number       Program_Type  \
0   San Francisco    CA     94114                    3  Medicare/Medicaid   
1          Orange    CA     9284
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
maxpop = pd.read_sql_query("""
  SELECT m.*, p.*
  FROM {} m JOIN population p ON m.Zip_Code = p.zipcode
  ORDER BY p.population DESC
  """.format(MYTABLE), conn)

answer = ['1194016923']
### END SOLUTION
```

```python
maxpop
```

Output:
```
index                                     Provider_Name         NPI  \
0     277                      CPH HOSPITAL MANAGEMENT, LLC  1194016923   
1     259          COUNTY OF LOS ANGELES AUDITOR CONTROLLER  1134143415   
2      94                     VERITAS HEALTH SERVICES, INC.  1962407460   
3     250               Jupiter Bellflower Doctors Hospital  1114021250   
4     169                       KAISER FOUNDATION HOSPITALS  1477608271   
..    ...                                               ...         ...   
82    231  GARDENS REGIONAL HOSPITAL AND MEDICAL CENTER INC  1043373053   
83    134               PASADENA HOSPITAL ASSOCIATION, LTD.  1407828429   
84    110            AVALON MEDICAL DEVELOPMENT CORPORATION  1346250347   
85     38                 UNIVERSITY OF SOUTHERN CALIFORNIA  1770728438   
86     79           University of California at Los Angeles  1902803315   

      CCN Medicaid_EP_Hospital_Type          Street_Address              City  \
0   50771      Acute C
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)

### BEGIN HIDDEN TESTS
assert(answer == ['1194016923'])
### END HIDDEN TESTS
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answer = pd.read_sql_query("""
  SELECT z.Medicaid_EP_Hospital_Type, AVG(z.Payments / p.population) as Avg_Pay_per_Capita
  FROM (
      SELECT m.Medicaid_EP_Hospital_Type, m.Zip_Code, sum(Total_payments) as Payments
      FROM {} m
      GROUP BY m.Medicaid_EP_Hospital_Type, m.Zip_Code
  ) z JOIN population p ON z.Zip_Code = p.zipcode
  GROUP BY
    z.Medicaid_EP_Hospital_type
  """.format(MYTABLE), conn)
### END SOLUTIONS

```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```

```python
answer
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals        17974.362796
1      Children's Hospitals          111.559033
```

## Now using Pandas instead of SQL

```python
### BEGIN SOLUTION
total_pay_per_zip = df.groupby(['Medicaid_EP_Hospital_Type','Zip_Code'])['Total_payments'].sum().reset_index()
total_pay_per_zip['Zip_Code'] = total_pay_per_zip['Zip_Code'].astype(str)

populations = pd.read_sql_query("SELECT * FROM population", conn)
populations['zipcode'] = populations['zipcode'].astype(str)

merged = total_pay_per_zip.merge(populations, how='left', left_on='Zip_Code', right_on='zipcode')

merged['Avg_Pay_per_Capita'] = merged['Total_payments'] / merged['population']

answer = merged.groupby('Medicaid_EP_Hospital_Type')['Avg_Pay_per_Capita'].mean().reset_index()

answer
### END SOLUTIONS

```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals        17974.362790
1      Children's Hospitals          111.559021
```