# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
from pathlib import Path
HOME = str(Path.home())

ALLERGIES_FILE="/data/allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def allergy_count(json_file):
    """(str)->int
    This function takes a file name and reads the JSON from that
    file.  It assumes the file structure contains a dictionary
    with a key of 'entry' whose value is a list of allergies.
    
    This function returns the length of that list of entries."""
    
    with open(json_file) as f:
        allergies = json.load(f)
    
    return len(allergies['entry'])

### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_count(json_file):
    with open(json_file) as f:
        allergies = json.load(f)
        
    return(len(allergies.get('entry')))
### END SOLUTION
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
with open(ALLERGIES_FILE) as f:
    allergies = json.load(f)
```

```python
entry = allergies.get('entry')[0]
```

```python
patient = entry.get('resource').get('patient').get('display')
print(patient)
```

Output:
```
Jason Argonaut
```

```python
### BEGIN SOLUTION

def patient_count(json_file):
    """(str)->int
    This function reads the JSON-structured information from the file
    and identifies how many different (unique) patients for which there
    are allergies.  This function returns the number of unique patients.
    """
    
    with open(json_file) as f:
        allergies = json.load(f)
    
    patients = set()
    
    for allergy in allergies.get('entry'):
        patient = allergy.get('resource').get('patient').get('display')
        patients.add(patient)

    return len(patients)

### END SOLUTION
```

```python
### BEGIN SOLUTION
def patient_count(json_file):
    patients = set()
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients.add(name)
        
    return patients
### END SOLUTION
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
### BEGIN SOLUTION

def allergy_per_patient(json_file):
    """(str)->dict
    """
    
    with open(json_file) as f:
        allergies = json.load(f)
    
    patient_allergies = {}
    
    for allergy in allergies.get('entry'):
        patient = allergy.get('resource').get('patient').get('display')
        if patient in patient_allergies:
            patient_allergies[patient] += 1
        else:
            patient_allergies[patient] = 1

    return patient_allergies
### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_per_patient(json_file):
    patients = {}
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients[name] = patients.setdefault(name,0) + 1
        
    return patients
### END SOLUTION
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
{
    "Paul Boal": 1,
    "Jason Argonaut": 3
}
```

Output:
```
{}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list(ALLERGIES_FILE) == output

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---
---

# Stretch (Extra) Problems

Work on either of the stretch problems below can earn you up to 25 free points toward the midterm assignment.  That is, if you complete one of these extra problems successfully, you can skip 1 of the problems that will appear on the midterm exam coming up next week.

The midterm will be distribute this Saturday 3/13.

This assignment is due on Sunday 3/14.  If you are trying for one of these extra problems Slack me, and I'll provide you feedback on how you did on these before end of day Monday 3/15.  That way you can choose what to complete on the midterm.


---
---

### STRETCH for March 2021 - For those looking for an additional challenge

As I've mentioned in class, CMS is now enforcing a rule around price transparency.  Every facility that take Medicare payments is required to publish a "machine readable" file with it's pricing infomration for a number of common procedures across all of the payers they work with.  There are two examples of such files in the `/data/` directory: `whiteriver.json` and `saline.xml`.

If you want to compare contracted prices across these two hospitals, you'll need to read in the information from both of those files into some kind of data structure, then merge the data together from those two files.  See what you can do.

See if you can create an output file that has the following fields:
* HOSPITAL
* PROCEDURE_CODE
* PAYER
* AMOUNT

If you choose to work on this, you may get stuck at some point and you won't know if you're _doing it right_. Make some assumptions. Document your questions in this notebook.



```
Procedure Code |  Description  |  Gross Charges  |  Aetna  |  QualChoice
```

```python
import json

with open('whiteriver.json') as f:
    whiteriver = json.load(f)
```

```python
wr_charges = {}
for charge in whiteriver.get('root').get('StandardCharges'):
    procedure = charge.get('ProcedureCode')
    description = charge.get('Description')
    gross_charges = float(charge.get('OutpatientGrossCharge'))
    aetna = float(charge.get('AETNA_Outpatient'))
    qualchoice = float(charge.get('QUALCHOICE_Outpatient'))
    wr_charges[('whiteriver',procedure)] = [description, gross_charges, aetna, qualchoice]
```

```python
wr_charges[('whiteriver','00170')]
```

Output:
```
['ANESTH PROCEDURE ON MOUTH', 9546.58, 8400.9904, 5346.0848]
```

```python
wr_charges.keys()
```

Output:
```
dict_keys([('whiteriver', '00170'), ('whiteriver', '00220'), ('whiteriver', '01996'), ('whiteriver', '10004'), ('whiteriver', '10005'), ('whiteriver', '10006'), ('whiteriver', '10007'), ('whiteriver', '10008'), ('whiteriver', '10009'), ('whiteriver', '10010'), ('whiteriver', '10011'), ('whiteriver', '10012'), ('whiteriver', '10021'), ('whiteriver', '10030'), ('whiteriver', '10060'), ('whiteriver', '10061'), ('whiteriver', '10120'), ('whiteriver', '10140'), ('whiteriver', '10160'), ('whiteriver', '11010'), ('whiteriver', '11011'), ('whiteriver', '11012'), ('whiteriver', '11042'), ('whiteriver', '11043'), ('whiteriver', '11044'), ('whiteriver', '11045'), ('whiteriver', '11046'), ('whiteriver', '11047'), ('whiteriver', '11055'), ('whiteriver', '11056'), ('whiteriver', '11057'), ('whiteriver', '11102'), ('whiteriver', '11103'), ('whiteriver', '11104'), ('whiteriver', '11105'), ('whiteriver', '11106'), ('whiteriver', '11107'), ('whiteriver', '11200'), ('whiteriver', '11300'), ('whiteriver',
```

```python
from lxml import objectify

with open('saline.xml') as f:
    saline = objectify.parse(f)

```

```python
xml=saline.getroot()
facility = xml.find('Facility')
outpatient = facility.findall('Patient')[2]
```

```python
sal_charges={}

for charge in outpatient.findall('Charge'):
    charge_type = charge.attrib['Type']
    for item in charge.findall('Item'):
        aetna = None
        qualchoice = None
        code = item.attrib['Code']
        description = item.find('Description').text
        gross_charges = float(item.find('GrossCharge').text)
        for contract in item.find('Contracts').findall('Contract'):
            payer = contract.attrib['Payer']
            charge = contract.attrib['Charge']
            if payer == "AETNA - (POS)":
                aetna = float(charge)
            elif payer == "QUALCHOICE LIFE AND HEALTH INSURANCE COMPANY INC - (Indemnity)":
                qualchoice = float(charge)
        # print(code, gross_charges, aetna, qualchoice)
        sal_charges[('saline',code)] = [description, gross_charges, aetna, qualchoice]
```

```python
sal_charges
```

Output:
```
{('saline', '10060'): ['DRAINAGE OF SKIN ABSCESS', 3406.7896, 129.19, 1274.0],
 ('saline', '10061'): ['DRAINAGE OF SKIN ABSCESS', 3057.258571, None, None],
 ('saline', '10081'): ['DRAINAGE OF PILONIDAL CYST', 10689.97, None, None],
 ('saline', '10120'): ['REMOVE FOREIGN BODY', 5229.0, None, None],
 ('saline', '10140'): ['DRAINAGE OF HEMATOMA/FLUID', 4236.7, None, None],
 ('saline', '10160'): ['PUNCTURE DRAINAGE OF LESION', 878.0, None, None],
 ('saline', '10180'): ['COMPLEX DRAINAGE WOUND', 6049.0, None, None],
 ('saline', '11010'): ['DEBRIDE SKIN AT FX SITE', 5871.153333, None, None],
 ('saline', '11011'): ['DEBRIDE SKIN MUSC AT FX SITE', 3615.5, None, None],
 ('saline', '11042'): ['DEB SUBQ TISSUE 20 SQ CM/<', 878.0, None, 836.0],
 ('saline', '11043'): ['DEB MUSC/FASCIA 20 SQ CM/<', 2087.509375, None, None],
 ('saline', '11044'): ['DEB BONE 20 SQ CM/<', 6192.24, None, None],
 ('saline', '11055'): ['TRIM SKIN LESION', 493.0, None, None],
 ('saline', '11056'): ['TRIM SKIN LESIONS 2 TO 
```

```python
for code,charge in wr_charges.items():
    saline = sal_charges.get(('saline',code[1]))
    if not saline is None:
        print(code,charge,saline)
```

Output:
```
('whiteriver', '10060') ['DRAINAGE OF SKIN ABSCESS', 11506.95, 10126.116, 6443.892] ['DRAINAGE OF SKIN ABSCESS', 3406.7896, 129.19, 1274.0]
('whiteriver', '10061') ['DRAINAGE OF SKIN ABSCESS', 472.0, 415.36, 264.32] ['DRAINAGE OF SKIN ABSCESS', 3057.258571, None, None]
('whiteriver', '10120') ['REMOVE FOREIGN BODY', 732.0, 644.16, 409.92] ['REMOVE FOREIGN BODY', 5229.0, None, None]
('whiteriver', '10140') ['DRAINAGE OF HEMATOMA/FLUID', 8004.01, 7043.5288, 4482.2456] ['DRAINAGE OF HEMATOMA/FLUID', 4236.7, None, None]
('whiteriver', '10160') ['PUNCTURE DRAINAGE OF LESION', 1826.0, 1606.88, 1022.56] ['PUNCTURE DRAINAGE OF LESION', 878.0, None, None]
('whiteriver', '11010') ['DEBRIDE SKIN AT FX SITE', 34515.83, 30373.9304, 19328.8648] ['DEBRIDE SKIN AT FX SITE', 5871.153333, None, None]
('whiteriver', '11011') ['DEBRIDE SKIN MUSC AT FX SIT', 1405.0, 1236.4, 786.8] ['DEBRIDE SKIN MUSC AT FX SITE', 3615.5, None, None]
('whiteriver', '11042') ['DEB SUBQ TISSUE 20 SQ CM/<', 8412.01, 7402.5688,
```

---
---

### STRETCH from March 2020 - For those looking for an additional challenge

The Coronavirus is creating quite the stir right now.  There are some sources suggesting that trends show it is going to be significantly more serious than SARS was back in the 2002 timeframe.  Here's one visualization trying to demonstrate that: https://www.reddit.com/r/China_Flu/comments/ev2b4v/i_updated_some_charts_comparing_this_outbreak/

Someone on Kaggle has generously already compiled a dataset based on information from Johns Hopkins about the Coronavirus outbreak.  https://www.kaggle.com/brendaso/2019-coronavirus-dataset-01212020-01262020  Create a Kaggle account, if you don't already have one.  Download this data set and then upload it to your Jupyter Home folder.  (The "up arrow" button is for uploading a file.)

Use Python's built-in `csv` module to read the data from this file and generate the following information: **what are the total confirmed cases in all of Mainland China as of the latest information in the data set?**  Some important things to note:
* Each entry for a given city has the **cumulative** number of cases.  So that column is not additive (it cannot be summed).  You'll have to find a way to filter your data for the last day for each city, then total those up.
* If you choose to parse the date column, you will want to lookup how to do that using Python's `datetime` module.  Especially the `strptime` function.  https://docs.python.org/3/library/datetime.html#strftime-strptime-behavior  Hint: you can parse a date string in the format 2/17/2020 using the code below.  This link will tell you what things like `%m` and `%Y` mean.

```
from datetime import datetime
d = datetime.strptime('2/17/2020', '%m/%d/%Y')
```

If you want to take this another step, **create a list of tuples that contain (observate date, total confirmed) totalled over all locations represented in the data**

```python

```

```python

```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

---

```python
%%bash
git pull
git add week06_assignment_2.ipynb
git commit -a -m "Submitting the week 6 programming assignment"
git push
```


---

If the message above says something like _Submitting the week 6 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```